<subject>
Assignment How-To
</subject>

<details>
    
**Do Not Talk During Quizzes or Exams**

Do not talk at all once an exam has begun. You may talk again once you leave the room after the quiz or exam. The one exception to this rule is if you need to ask the exam proctor a question. Any talking during a quiz or exam qualifies the student for disciplinary action.

**Naming Conventions**

When naming your files for upload, you must follow the format below:

    <uni>_<assignment>_<details [optional]>.<extension>

For example, if I needed to hand in HW 0, any of the below formats would be sufficient for a file name:

- pl2648_hw0.ipynb
- pl2648_hw0.txt
- pl2648_hw0.sh
- pl2648_hw0_all_in_one.txt
- Pl2648_hw0_bash_program.sh
- Pl2648_quiz1.ipynb

This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

Courseworks may rename your file to something like `pl2648_hw0-1.ipynb` if you resubmit your assignment. This is perfectly fine.

**What Format To Submit In**

Most homework and quizzes are in Jupyter notebooks. Unless specified otherwise, please download your work as an `.ipynb` file from your local machine and upload it to courseworks.


**Grading**

Possible points on late assignments are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors. 

Because grading is automated, please delete (or comment out) the `raise NotImplmeneted` code before attempting a problem. 

Empty un-editable cells in an assignment are there for a purpose. They will be filled with tests by the automatic grader. Please do not attempt to remove them.

**Getting Help**

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. You can not post screenshots of your current work to Piazza or other tools used for getting help.

If you need to reach out to a CA for help, please do so via Piazza and not via email. Answers given via Piazza will help you as well as other students. Thus, emails will always have a lower priority for response than Piazza questions. If you do email the CA, please make a note of what section you are in. This helps us identify you in courseworks faster. 

Finally, if you do not get a repsonse from a CA within 48 hours, you may email the professor.

**Multiple Choice**

If the question is multiple choice, you will be given several options to choose from and your function will need to return **one** of those options **verbatim** as a string.

For example: 

Which of the following animals bark?

- dogs
- cats
- fish
- trees

A correct answer would be structured in the following way:

```python
def question_animals():
    return 'dogs'
```

You answer will be stripped of left and right white space and lowercased before comparison to the correct answer during grading.
</details>

**Q**: Write a command line script that opens a file, reads its lines into a buffer, sorts those lines alphabetically, and prints each line back out to standard out in sorted order. The script should take a single argument passed in at the command line and that argument should be the path to a file. The script should make use of the ArgumentParser class. Empty lines do not need to be preserved.

Example:

Assume `/path/to/file` holds the following contents.

```
b
c
a
```

Then you script should operate as shown.


```bash
$ python script.py /path/to/file
a
b
c
```

`get_sort_script` should return a multiline string that contains the contents of the script `script.py`. It will most likely be easiest to build the script on a VM or your own machine first and then copy it into place in the function below.

The script should also make use of the `if __name__ == '__main__':` pattern.

Do not use a shebang line ('#!..') in this particular script. 

Sorry Windows users, but you may need to run this notebook in a Linux or MacOS environment to get the test below to work correctly. Your homework will be tested in Linux or MacOS environment so if you are confident your code is correct, no need to worry if your tests do not pass on your Windows machine. 

[3 points]

<!-- grade_id: cell-e25240502b1eb2a9, points: 0, grade: false, solution: true -->
```python
def get_sort_script() -> str:
    return """
def main(file):
    rows=file.readlines()
    rows2 = [line for line in rows if line.strip() != ""]
    sorted_rows = sorted(rows2)
    for row in sorted_rows:
        line=row.strip()
        print(line)

if __name__ == '__main__':
    import argparse

    parser = argparse.ArgumentParser()
    parser.add_argument('file', type=argparse.FileType('r'))
    args = parser.parse_args()

    main(args.file)
    """

get_sort_script()
```

Output:
```
'\ndef main(file):\n    rows=file.readlines()\n    rows2 = [line for line in rows if line.strip() != ""]\n    sorted_rows = sorted(rows2)\n    for row in sorted_rows:\n        line=row.strip()\n        print(line)\n\nif __name__ == \'__main__\':\n    import argparse\n\n    parser = argparse.ArgumentParser()\n    parser.add_argument(\'file\', type=argparse.FileType(\'r\'))\n    args = parser.parse_args()\n\n    main(args.file)\n    '
```

<!-- grade_id: cell-6a0af5b6d334960e, points: 0.5, grade: true, solution: false -->
```python
assert 'ArgumentParser' in get_sort_script(), 'No use of ArgumentParser'
```

<!-- grade_id: cell-d3214f6d919edfa0, points: 0.5, grade: true, solution: false -->
```python
assert "if __name__ == '__main__':" in get_sort_script(), 'No use of __name__ == __main__'
```

```python
import subprocess
import sys
import tempfile


def run_script_on_data(data):
    with tempfile.NamedTemporaryFile() as unsorted_file:
        unsorted_file.write(data.encode())
        unsorted_file.flush()

        with tempfile.NamedTemporaryFile() as script:
            script.write(get_sort_script().encode())
            script.flush()

            out = subprocess.run([sys.executable, script.name, unsorted_file.name], capture_output=True)
            
            if out.stderr:
                raise RuntimeError(out.stderr.decode())

            return out.stdout.decode()
```

<!-- grade_id: cell-3c009204c637356d, points: 2, grade: true, solution: false -->
```python
import random

data = list(map(str, (random.randint(0, 100) for _ in range(10))))

result = run_script_on_data('\n'.join(data)).split()

data.sort()

assert result == data, (result, '!=', data)
```

Output:
```
<ipython-input-4-f1b8312c5eb8> in run_script_on_data(data)
     13             script.flush()
     14 
---> 15             out = subprocess.run([sys.executable, script.name, unsorted_file.name], capture_output=True)
     16 
     17             if out.stderr:

~\Anaconda3\lib\subprocess.py in run(input, timeout, check, *popenargs, **kwargs)
    401         kwargs['stdin'] = PIPE
    402 
--> 403     with Popen(*popenargs, **kwargs) as process:
    404         try:
    405             stdout, stderr = process.communicate(input, timeout=timeout)

TypeError: __init__() got an unexpected keyword argument 'capture_output'
```

```python
import random

data = list(map(str, (random.randint(0, 100) for _ in range(10))))


result = run_script_on_data('\n'.join(data)).split()

#data.sort()

```

Output:
```
<ipython-input-4-f1b8312c5eb8> in run_script_on_data(data)
     13             script.flush()
     14 
---> 15             out = subprocess.run([sys.executable, script.name, unsorted_file.name], capture_output=True)
     16 
     17             if out.stderr:

~\Anaconda3\lib\subprocess.py in run(input, timeout, check, *popenargs, **kwargs)
    401         kwargs['stdin'] = PIPE
    402 
--> 403     with Popen(*popenargs, **kwargs) as process:
    404         try:
    405             stdout, stderr = process.communicate(input, timeout=timeout)

TypeError: __init__() got an unexpected keyword argument 'capture_output'
```