<subject>
Assignment How-To
</subject>

<details>
    
**Do Not Talk During Quizzes or Exams**

Do not talk at all once an exam has begun. You may talk again once you leave the room after the quiz or exam. The one exception to this rule is if you need to ask the exam proctor a question. Any talking during a quiz or exam qualifies the student for disciplinary action.

**Naming Conventions**

When naming your files for upload, you must follow the format below:

    <uni>_<assignment>_<details [optional]>.<extension>

For example, if I needed to hand in HW 0, any of the below formats would be sufficient for a file name:

- pl2648_hw0.ipynb
- pl2648_hw0.txt
- pl2648_hw0.sh
- pl2648_hw0_all_in_one.txt
- Pl2648_hw0_bash_program.sh
- Pl2648_quiz1.ipynb

This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

Courseworks may rename your file to something like `pl2648_hw0-1.ipynb` if you resubmit your assignment. This is perfectly fine.

**What Format To Submit In**

Most homework and quizzes are in Jupyter notebooks. Unless specified otherwise, please download your work as an `.ipynb` file from your local machine and upload it to courseworks.


**Grading**

Possible points on late assignments are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors. 

Because grading is automated, please delete (or comment out) the `raise NotImplmeneted` code before attempting a problem. 

Empty un-editable cells in an assignment are there for a purpose. They will be filled with tests by the automatic grader. Please do not attempt to remove them.

**Getting Help**

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. You can not post screenshots of your current work to Piazza or other tools used for getting help.

If you need to reach out to a CA for help, please do so via Piazza and not via email. Answers given via Piazza will help you as well as other students. Thus, emails will always have a lower priority for response than Piazza questions. If you do email the CA, please make a note of what section you are in. This helps us identify you in courseworks faster. 

Finally, if you do not get a repsonse from a CA within 48 hours, you may email the professor.

**Multiple Choice**

If the question is multiple choice, you will be given several options to choose from and your function will need to return **one** of those options **verbatim** as a string.

For example: 

Which of the following animals bark?

- dogs
- cats
- fish
- trees

A correct answer would be structured in the following way:

```python
def question_animals():
    return 'dogs'
```

You answer will be stripped of left and right white space and lowercased before comparison to the correct answer during grading.
</details>

**Q**: Write a function `get_cubes` that builds a new list via a list comprehension. The list comprehension should contain in sorted order, all cubes for numbers between 0 a given number `x` if the number is divisible by 3.

```python
>>> get_cubes(10)
[0, 27, 216, 729]
>>> get_cubes(0)
[]
>>> get_cubes(9)
[0, 27, 216]
```

[1 point]

<!-- grade_id: cell-70c782eb764fa9f9, points: 0, grade: false, solution: true -->
```python
def get_cubes(x):
    return [x**3 for x in range(x) if x % 3 == 0]
get_cubes(9)
```

Output:
```
[0, 27, 216]
```

<!-- grade_id: cell-a6c8d30e831fe94e, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert get_cubes(9) == [0, 27, 216]
### BEGIN TESTS
```

<!-- grade_id: cell-f25ee1a76203c90e, points: 0, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-7f2062eecdd55c4f, points: 0.5, grade: true, solution: false -->
```python

```

**Q**: The function call `random.choice(string.ascii_lowercase)` returns a random lowercased ASCII letter. Create a function `get_count` that returns the number of unique letters returned by `x` sequential calls of `random.choice(string.ascii_lowercase)`. Your function should use a set comprehension.

```python
>>> get_count(10)
8
>>> get_count(10)
7
>>> get_count(10)
9
>>> get_count(7)
5
```

[1 point]

<!-- grade_id: cell-00e884959ab1fddd, points: 0, grade: false, solution: true -->
```python
import string
import random

def get_count(x):
    l={random.choice(string.ascii_lowercase) for x in range(x)}
    return len(l)

get_count(10)
```

Output:
```
6
```

<!-- grade_id: cell-b57633eea1568b62, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert get_count(1) == 1
assert get_count(5) <= 5
assert get_count(10) <= 10
### END TESTS
```

<!-- grade_id: cell-06fffca9a7420049, points: 0.5, grade: true, solution: false -->
```python

```

**Q**: Create a function `get_debits` that takes a dictionary and returns a new dictionary were each key in the given dictionary is prefixed with "debit\_" and each amount's value is multiplied by -1. The function should make use of a dictionary comprehension.


```python
>>> d = {'helicopter': 430, 'hanger': 75}
>>> get_debits(d)
{'debit_hanger': -75, 'debit_helicopter': -430}
```

[1 point]

<!-- grade_id: cell-225b03fcb668cf76, points: 0, grade: false, solution: true -->
```python
def get_debits(d):
    dict_variable = {"debit_" + key:value*-1 for (key,value) in d.items()}
    return dict_variable
        
d = {'helicopter': 430, 'hanger': 75}
get_debits(d)
```

Output:
```
{'debit_helicopter': -430, 'debit_hanger': -75}
```

<!-- grade_id: cell-640e5124f8e53cbc, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert get_debits({'helicopter': 430, 'hanger': 75}) == {'debit_helicopter': -430, 'debit_hanger': -75}
### END TESTS
```

<!-- grade_id: cell-4ad3da6ea4b02b76, points: 0.5, grade: true, solution: false -->
```python

```

**Q**: Create a Ceaser cipher. Implement a function `ceaser_encrypt` that takes some text (`plaintext`) and returns ciphertext. To compute the ciphertext, iterate through each letter in the plain text and replace it with the letter that is 13 characters forward in the ASCII alphabet. For letters that are too far towards the end of the alphabet, start back at the beginning of the alphabet and continue counting up to 13. The function should convert `plaintext` to lower case before operating on the text. All punctuation, numbers, and spaces should be ignored by the cipher.

For example:

```python
>>> def ceaser_encrypt(plaintext):
...     # implement me
...     pass
... 
>>> ceaser_encrypt('You rock!')
'lbh ebpx!'
```
[1 point]

For fun, see what the result is if you pass text encrypted by `ceaser_encrypt` back into `ceaser_encrypt`.

<!-- grade_id: cell-aad5279e1353f544, points: 0, grade: false, solution: true -->
```python
import string

def caeser_encrypt(text, step, alphabets):

    def shift(alphabet):
        return alphabet[step:] + alphabet[:step]

    shifted_alphabets = tuple(map(shift, alphabets))
    joined_aphabets = ''.join(alphabets)
    joined_shifted_alphabets = ''.join(shifted_alphabets)
    table = str.maketrans(joined_aphabets, joined_shifted_alphabets)
    return text.translate(table)

alphabets = (string.ascii_lowercase, string.ascii_uppercase, string.digits)
#caeser_encrypt('You rock!', step=13, alphabets=alphabets)

ceaser_encrypt('clguba vf terng')


ceaser_encrypt('ceaser_encrypt')
```

Output:
```
'prnfre_rapelcg'
```

<!-- grade_id: cell-2b614b795bd55609, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert ceaser_encrypt('clguba vf terng') == 'python is great'
### BEGIN TESTS
```

<!-- grade_id: cell-e0ac874facd027c9, points: 0.5, grade: true, solution: false -->
```python

```

**Q**: Write a function named `factorial` to find the factorial for a given number. The function should return an integer. You can assume the argument given is non-negative.

```
>>> factorial(3)
6
```

[1 point]

<!-- grade_id: cell-e103feeaf57201f9, points: 0, grade: false, solution: true -->
```python
import math 

def factorial(n):
       return (math.factorial(n)) 
        
factorial(3)
```

Output:
```
6
```

<!-- grade_id: cell-9f0ec1966bdd4739, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
assert factorial(3) == 6
### BEGIN TESTS
```

<!-- grade_id: cell-84c815a70530d4da, points: 0.5, grade: true, solution: false -->
```python

```

**Extra Practice**: Write a script that acts like a stopwatch. You should be prompted to press enter to start the watch. Once the watch is running, you can create a lap by pressing enter again. Every time "ENTER" is pressed, the current lap time should be printed. Lap time is time that has elapsed since the last press of the "ENTER" button.

```python
>>> def timer():
...     # implement me
...     pass
>>> timer()
PRESS ENTER >
LAP: 0.0
PRESS ENTER >
LAP: 2.182952880859375
PRESS ENTER >
LAP: 2.125649929046631
PRESS ENTER >
```

[0 points]

<!-- grade_id: cell-b29f023d4c702a2f, points: 0, grade: false, solution: true -->
```python
import time

def timer():
    # YOUR CODE HERE
    raise NotImplementedError()
    
# timer()  # uncomment me to test me out
```